# yamllint disable rule:line-length
---
rule_files:
  - ../prometheus_alerts.yaml

tests:
  - interval: 1m
    input_series:
      - series: 'envoy_cluster_upstream_rq_xx{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster", envoy_response_code_class="2"}'
        values: "1+1000x10"
      - series: 'envoy_cluster_upstream_rq_xx{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster", envoy_response_code_class="4"}'
        values: "1+500x10"
      - series: 'envoy_cluster_upstream_rq_total{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster"}'
        values: "2+1500x10"
    alert_rule_test:
      - eval_time: 10m
        alertname: EnvoyUpstreamHighHttp4xxErrorRate
        exp_alerts:
          - exp_labels:
              namespace: default
              envoy_cluster_name: test-cluster
              severity: info
            exp_annotations:
              summary: "Envoy upstream high HTTP 4xx error rate."
              description: "More than 5% HTTP requests with status 4xx for cluster test-cluster in default the past 5m."
              dashboard_url: "https://grafana.com/d/envoy-upstream-skj2/envoy-upstream?var-namespace=default&var-envoy_cluster_name=test-cluster"
  - interval: 1m
    input_series:
      - series: 'envoy_cluster_upstream_rq_xx{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster", envoy_response_code_class="2"}'
        values: "1+1000x10"
      - series: 'envoy_cluster_upstream_rq_xx{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster", envoy_response_code_class="5"}'
        values: "1+500x10"
      - series: 'envoy_cluster_upstream_rq_total{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster"}'
        values: "2+1500x10"
    alert_rule_test:
      - eval_time: 10m
        alertname: EnvoyUpstreamHighHttp5xxErrorRate
        exp_alerts:
          - exp_labels:
              namespace: default
              envoy_cluster_name: test-cluster
              severity: critical
            exp_annotations:
              summary: "Envoy upstream high HTTP 5xx error rate."
              description: "More than 5% HTTP requests with status 5xx for cluster test-cluster in default the past 5m."
              dashboard_url: "https://grafana.com/d/envoy-upstream-skj2/envoy-upstream?var-namespace=default&var-envoy_cluster_name=test-cluster"
  - interval: 1m
    input_series:
      - series: 'envoy_cluster_circuit_breakers_default_rq_open{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster"}'
        values: "0 0 0 0 0 1+0x10"
    alert_rule_test:
      - eval_time: 10m
        alertname: EnvoyCircuitBreakerOpen
        exp_alerts:
          - exp_labels:
              namespace: default
              envoy_cluster_name: test-cluster
              severity: warning
            exp_annotations:
              summary: "Envoy circuit breaker is open."
              description: "Circuit breaker is open for cluster test-cluster in default for the past 5m."
              dashboard_url: "https://grafana.com/d/envoy-upstream-skj2/envoy-upstream?var-namespace=default&var-envoy_cluster_name=test-cluster"
  - interval: 1m
    input_series:
      - series: 'envoy_cluster_upstream_cx_connect_fail{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster"}'
        values: "1+6100x20"
    alert_rule_test:
      - eval_time: 20m
        alertname: EnvoyUpstreamConnectionFailures
        exp_alerts:
          - exp_labels:
              namespace: default
              envoy_cluster_name: test-cluster
              severity: warning
            exp_annotations:
              summary: "Envoy upstream connection failures detected."
              description: "More than 100 connection failures for cluster test-cluster in default the past 5m."
              dashboard_url: "https://grafana.com/d/envoy-upstream-skj2/envoy-upstream?var-namespace=default&var-envoy_cluster_name=test-cluster"
  - interval: 1m
    input_series:
      - series: 'envoy_cluster_membership_total{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster"}'
        values: "10+0x10"
      - series: 'envoy_cluster_membership_healthy{job="envoy-metrics", namespace="default", pod="envoy-test", envoy_cluster_name="test-cluster"}'
        values: "6+0x10"
    alert_rule_test:
      - eval_time: 10m
        alertname: EnvoyUpstreamUnhealthyHosts
        exp_alerts:
          - exp_labels:
              namespace: default
              envoy_cluster_name: test-cluster
              severity: warning
            exp_annotations:
              summary: "Envoy upstream has unhealthy hosts."
              description: "More than 33% of hosts are unhealthy for cluster test-cluster in default for the past 5m."
              dashboard_url: "https://grafana.com/d/envoy-upstream-skj2/envoy-upstream?var-namespace=default&var-envoy_cluster_name=test-cluster"
  - interval: 1m
    input_series:
      - series: 'xds_snapshot_update_total{job="envoy-gateway", namespace="envoy-gateway-system", status="error", nodeID="node-1"}'
        values: "0+1x10"
      - series: 'xds_snapshot_update_total{job="envoy-gateway", namespace="envoy-gateway-system", status="success", nodeID="node-1"}'
        values: "0+100x10"
    alert_rule_test:
      - eval_time: 10m
        alertname: EnvoyXDSUpdateFailed
        exp_alerts:
          - exp_labels:
              namespace: envoy-gateway-system
              status: error
              nodeID: node-1
              severity: warning
            exp_annotations:
              summary: "Envoy Gateway XDS snapshot update failed."
              description: "XDS snapshot update failed for node node-1 in envoy-gateway-system with status error the past 5m."
              dashboard_url: "https://grafana.com/d/envoy-gateway-overview-skj2/envoy-gateway-overview?var-namespace=envoy-gateway-system"
